```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 2: Pandas and Data Visualization
# Due Friday 1/23 at 11:59pm

```python
# Import me!
import numpy as np
import pandas as pd
from matplotlib import pyplot as plt
```

## Point Breakdown

Question | Sub | Points | Manual
----|----|----|----
Q0|a|2|no|
Q0|b|1|no|
Q1|aa|2|no|
Q1|ab|5|no|
Q1|ac|5|no|
Q1|b|10|no|
Q1|c|15|yes|
Q1|d|5|yes|
Q1|e|15|yes|
Q1|f|15|yes|
Q1|g|5|yes|
Q1|h|15|yes|
Q1|i|5|yes|

---

## Question 0: Warm Up - Creating a Pandas DataFrame 

We want to create a Pandas DataFrame from the following information:

> The common name for *Yucca filamentosa* is Yucca, and it is not native to Oregon. Two plants that are native to Oregon are *Neottia cordata*, known as Heart-leaved twayblade, and *Festuca idahoensis*, known as Idaho fescue.
  
### Question 0a (2 pts):
First make a dictionary from this information with three keys: `'latin_name'`, `'common_name'`, and `'oregon_native'`. The key `'oregon_native'` should have values `'yes'` or `'no'`.

**Hint**: Copy/paste the names into your dictionary to avoid spelling issues. For the latin names, do not copy the asterisks. 

#### Answer 0a

```python
flora_dict = ...
```

```python
grader.check("q0a")
```

### Question 0b (1 pt):

Now convert `flora_dict` into a 3x3 Pandas DataFrame.

#### Answer 0b

```python
flora_df = ...
```

```python
grader.check("q0b")
```

---
## Question 1: 

Bike sharing systems are a new generation of traditional bike rentals where the process of signing up, renting and returning is automated. We will be analyzing bike sharing data from Washington D.C. in 2011 and 2012.

In this assignment, you will perform tasks to clean, visualize, and explore the bike sharing data. 

After completing this assignment, you should be comfortable with:

* Wrangling data for analysis
* Using exploratory data analysis to investigate the data
* Making informative visualizations and controlling plot parameters
* Grouping data with `groupby` and `agg`


### Question 1aa (2pts)
#### Loading Bike Sharing Data

Read the csv file `bikeshare.txt` into your notebook as a Pandas dataframe named `bikeshare_raw`. The resulting dataframe will have a series of variable names (columns). Below is a description of each variable in original file.

Variable       | Description
-------------- | ------------------------------------------------------------------
instant | record index
datetime | date (string)
season | 1. spring <br> 2. summer <br> 3. fall <br> 4. winter
yr | year (0: 2011, 1: 2012)
mnth | month ( 1 to 12)
hr | hour (0 to 23)
holiday | whether day is holiday or not
weekday | day of the week
workingday | if day is neither weekend nor holiday
weathersit | 1. clear or partly cloudy <br> 2. mist and clouds <br> 3. light snow or rain <br> 4. heavy rain or snow
temp | normalized temperature in Celsius (divided by 41)
atemp | normalized "feels-like" temperature in Celsius (divided by 50)
hum | normalized percent humidity (divided by 100)
windspeed| normalized wind speed (divided by 67)
casual | count of casual bike users
registered | count of registered bike users
cnt | count of total rental bikes including casual and registered  


```python
# Load data into a Pandas Data Frame...
bikeshare_raw = ...
```

```python
grader.check("q1aa")
```

#### Take a look at the dataframe

```python
bikeshare_raw.head()
```

### Question 1ab (5pts)
#### Making replacements
The `replace` [method](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.replace.html) of Pandas dataframes allows us to use a dictionary to replace existing values with alternatives. 
  
Create a replacement dictionary and use it within the `.replace()` method to make the following changes to bikeshare_raw. 

 1. `weekday` numbers (0,1,2,3,4,5,6) to three-letter string labels (`Sun`, `Mon`, `Tue`, `Wed`, `Thu`, `Fri`, and `Sat`).  
   
 2. Change the `holiday` and `workingday` labels of  `1` and `0` to `yes` and `no` respectively. 

Assign your replacement dictionary to the variable `replacement_dict`.  
Don't edit bikeshare_raw in-place, instead assign the output dataframe to the variable `bikeshare`.

```python
# Assign your replacement dictionary to the variable replacement_dict
replacement_dict = ...

# Use .replace() on bikeshare_raw (do not do this in-place)
# assign the output dataframe to the variable bikeshare
bikeshare = ...
bikeshare.head()
```

```python
grader.check("q1ab")
```

### Question 1ac (5 pts)
#### Converting hourly rider data to daily counts 
The original dataset has entries at the hourly level. However, for initial analysis, daily statistics will be most useful. We can do this using the `groupby` method introduced in class.

Use `groupby` to construct a Pandas Series named `mean_temp_by_day`. The values should be mean temperature for each unique date. The index should be the datetime value.

```python
mean_temp_by_day = ...
```

```python
grader.check("q1ac")
```

### Question 1b (10 pts)

Construct a DataFrame named `daily_counts` grouped by `datetime` with the following columns:
* `casual`: total number of casual riders for each day
* `registered`: total number of registered riders for each day
* `workingday`: whether that day is a working day or not (`yes` or `no`)
* `mean_temp` : average temperature across all hourly readings from the day
* `max_temp` : maximum temperature from the day
* `min_temp` : minimum temperature from the day

Finally, `datetime` should end up as its own column, not as the index - use `reset_index` for making dates a column, not an index.

**Hint**: Use `groupby` and `agg`. If you use the capability to do different aggregations by column, you can do this task with a single call to `groupby` and `agg`. For the `workingday` column we can take any of the values since we are grouping by the day, thus the value will be the same within each group. Take a look at the `'first'` or `'last'` aggregation functions.

```python
# Use groupby() and agg() to create the dataframe `daily_counts`
daily_counts = ...

# Reset the index of your daily_counts dataframe 
daily_counts = daily_counts.reset_index()
```

```python
daily_counts.head()
```

```python
grader.check("q1b")
```

### Question 1c (15pts)
#### Exploratory plotting of casual and registered ridership over time
Make a scatter plot of daily casual and registered ridership over time that adheres to the following specifications:

   1. Date is the horizontal axis. You can use the `datetime` column directly for these values. Some code is included to help clean up the label formatting.
   2. Number of Daily Riders is on the vertical axis.

   3. Include a legend, title, and axis labels.    

#### Notes on plot style:
To copy our style in the example plot, use the following arguments in your ax1c.scatter() calls.
- `facecolors='none'` (creates hollow points)
- `edgecolors='c'`  (plot cyan edge color; you can assign color you want for distinct points)

Your plot should look something like this:

<img src="casual_and_registered_riders_daily.png">

*This question will be manually graded. The checks below are there to make sure your plot has all of the required elements*

<!-- BEGIN QUESTION -->



```python
# This code creates your figure and axis
fig,ax1c=plt.subplots(figsize=(8,8))

# Add elements to ax1c e.x. ax1c.scatter(...)
...

# This code formats your x-tick labels
plt.xticks(np.arange(0,len(daily_counts), 60), rotation = 45)
fig.autofmt_xdate()
```

```python
grader.check("q1c")
```

<!-- END QUESTION -->

### Question 1d (5pts)
- What does the above scatter plot suggest about possible seasonal trends in casual and registered ridership? 
- Apart from seasonal trends, what other sources of variability might contribute to daily differences in casual and registered ridership? 
- What does the plot tell us about the relative numbers of casual and registered riders?

_Type your answer here, replacing this text._

### Question 1e (15 pts)
Make a scatter plot of registered riders vs casual riders for each day, following these specifications:

1. x axis is the # of casual riders
2. y axis is the # of registered riders
3. Set the aspect ratio of the plot to 1 so that the horizontal and vertical axes have the same scale. 
4. Add a line with slope of 1 and intercept 0,  which corresponds to casual and registered ridership being equal. 
5. Include axis labels and a title. 

**Hint**: The `axis` method for setting aspect is `set_aspect(...)` [(documentation)](https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.set_aspect.html).

Your plot should look something like this:

<img src="casual_vs_registered_ridership_compared_1.png">

*This question will be manually graded. The checks below are there to make sure your plot has all of the required elements*

<!-- BEGIN QUESTION -->



```python
# This creates a new figure and sets its size
fig, ax1e = plt.subplots(figsize=(8,8))

# Add elements to ax1e e.x. ax1e.scatter(...)
...
```

```python
grader.check("q1e")
```

<!-- END QUESTION -->

### Extra Question (0 pts)
Are there certain days with more casual ridership than registered? <br>
You can read this off the graph rather than perform analysis. How many days are there?

_Type your answer here, replacing this text._

### Question 1f (15 pts)
Modify the above scatter plot so that the points plotted are divided into two colors. One color should denote working days (values of `workingday` with value `'yes'`) and another color denotes non-working days (weekends and holidays; values of `workingday` with value `'no'`). </br>

Add a legend, title, axis labels, and a unit line (with slope 1 that goes through origin).

Your plot should look something like this:

<img src="casual_vs_registered_ridership_compared_2.png">

*This question will be manually graded. The checks below are there to make sure your plot has all of the required elements*

### Answer 1f

<!-- BEGIN QUESTION -->



```python
# This creates a new figure and sets its size
fig, ax1f = plt.subplots(figsize=(8,8))

# Add elements to ax1f e.x. ax1f.scatter(...)
...
```

```python
grader.check("q1f")
```

<!-- END QUESTION -->

### Question 1g (5 pts)
On days with higher numbers of casual riders, do there also seem to be higher amounts of registered riders?<br>
Does the value of `working day` seem to explain the ratio between casual and registered riders for that day? Explain.  <br>


### Answer 1g: 

_Type your answer here, replacing this text._

### Question 1h (15 pts)
#### Testing the influence of temperature on casual and registered ridership
Make a scatter plot that shows the relationship between daily mean temperature and casual ridership on working days and nonworking days. Plot individual days as scatter points, with `daily mean temperature` represented on the horizontal axis and `casual ridership` on the vertical axis. Make `working` and `nonworking` days separate colors to differentiate them.</br>

Add legend, title, and axis labels  

Your plot should look something like this:

<img src="casual_ridership_vs_mean_temp.png">

*This question will be manually graded. The checks below are there to make sure your plot has all of the required elements*

### Answer 1h

<!-- BEGIN QUESTION -->



```python
# This creates a new figure and sets its size
fig, ax1h = plt.subplots(figsize=(8,8))

# Add elements to ax1h e.x. ax1h.scatter(...)
...
```

```python
grader.check("q1h")
```

<!-- END QUESTION -->

### Question 1i (5 pts)

Please answer each of the following questions in the space allowed below:
1. Does casual ridership tend to be higher on working or non-working days?
2. Broadly speaking, regardless of whether it is a working day, does the amount of casual ridership seem dependent on mean-temperature? How so?
3. Is the difference in casual ridership between working and non-working days higher on warmer or colder days?

### Answer 1i

_Type your answer here, replacing this text._

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Save your notebook and then run the last cell. Download the generated zip file and upload it to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q0a (points: 2)
```
>>> assert len(flora_dict) == 3
```
failure_message: the length of flora_dict is incorrect
```
>>> assert set(flora_dict.keys()) == set(['latin_name', 'common_name', 'oregon_native'])
```
failure_message: the keys of flora_dict are incorrect
```
>>> flora_df_rowise_test = [set(row) for row in pd.DataFrame(flora_dict).values]
>>> assert set(['Yucca', 'Yucca filamentosa', 'no']) in flora_df_rowise_test
>>> assert set(['Heart-leaved twayblade', 'Neottia cordata', 'yes']) in flora_df_rowise_test
>>> assert set(['Festuca idahoensis', 'Idaho fescue', 'yes']) in flora_df_rowise_test
```
failure_message: the values of flora_dict are incorrect


Tests for q0b (points: 1)
```
>>> assert flora_df.shape == (3, 3)
```
failure_message: the shape of flora_df is incorrect
```
>>> assert set(flora_df.columns) == set(['common_name', 'latin_name', 'oregon_native'])
```
failure_message: the column names of flora_df are incorrect
```
>>> flora_df_rowise_test = [set(row) for row in flora_df.values]
>>> assert set(['Yucca', 'Yucca filamentosa', 'no']) in flora_df_rowise_test
>>> assert set(['Heart-leaved twayblade', 'Neottia cordata', 'yes']) in flora_df_rowise_test
>>> assert set(['Festuca idahoensis', 'Idaho fescue', 'yes']) in flora_df_rowise_test
```
failure_message: the values of flora_df are incorrect


Tests for q1aa (points: 2)
```
>>> assert type(bikeshare_raw) == pd.DataFrame
```
failure_message: Check that bikeshare_raw is a pandas dataframe
```
>>> assert bikeshare_raw.shape == (17379, 17)
```
failure_message: The shape of bikeshare_raw should be 17379 rows and 17 columns


Tests for q1ab (points: 5)
```
>>> assert isinstance(replacement_dict, dict)
```
failure_message: make sure replacement_dict is a dictionary
```
>>> assert set(replacement_dict.keys()) == set(bikeshare_raw.columns[[6, 7, 8]])
```
failure_message: make sure the keys of replacement dict are the names of the columns specified in the instructions
```
>>> assert set(replacement_dict['workingday'].items()) == set({(0, 'no'), (1, 'yes')})
```
failure_message: for the workingday key, make sure you replace 0 with no and 1 with yes
```
>>> assert set(replacement_dict['holiday'].items()) == set({(0, 'no'), (1, 'yes')})
```
failure_message: for the holiday key, make sure you replace 0 with no and 1 with yes
```
>>> assert set(replacement_dict['weekday'].items()) == set({(0, 'Sun'), (1, 'Mon'), (2, 'Tue'), (3, 'Wed'), (4, 'Thu'), (5, 'Fri'), (6, 'Sat')})
```
failure_message: for the weekday key, make sure you replace 0-6 with the corresponding weekday abbreviations specified in the question
```
>>> assert isinstance(bikeshare, pd.DataFrame)
```
failure_message: make sure bikeshare is a pandas dataframe
```
>>> assert bikeshare.shape == (17379, 17)
```
failure_message: the shape of bikeshare should still be 17379 rows and 17 columns
```
>>> from io import StringIO
>>> bikeshare_raw_test = pd.read_csv('bikeshare.txt')
>>> bikeshare_replace = pd.read_csv(StringIO('value,holiday,workingday,weekday\n0,no,no,Sun\n1,yes,yes,Mon\n2,,,Tue\n3,,,Wed\n4,,,Thu\n5,,,Fri\n6,,,Sat'))
>>> assert bikeshare_raw_test.replace(bikeshare_replace).equals(bikeshare)
```
failure_message: check that you are assigning your replaced dataframe to bikeshare


Tests for q1ac (points: 5)
```
>>> assert isinstance(mean_temp_by_day, pd.Series)
```
failure_message: mean_temp_by_day should be a pandas series
```
>>> assert mean_temp_by_day.shape == (731,)
```
failure_message: mean_temp_by_day should have 731 rows
```
>>> assert mean_temp_by_day.index.name == 'datetime'
```
failure_message: check that you are grouping by datetime
```
>>> def stat_summary_check(data):
...     min_val = np.min(data)
...     q25 = np.percentile(data, 25)
...     median = np.median(data)
...     q75 = np.percentile(data, 75)
...     max_val = np.max(data)
...     mean = np.mean(data)
...     std_dev = np.std(data, ddof=1)
...     return tuple((float(round(val, 2)) for val in (min_val, q25, median, q75, max_val, mean, std_dev)))
>>> check_tuple = (0.06, 0.34, 0.5, 0.66, 0.86, 0.5, 0.18)
>>> assert np.allclose(stat_summary_check(mean_temp_by_day), check_tuple, 0.001)
```
failure_message: check that you are getting the mean of the temp column for each day


Tests for q1b (points: 10)
```
>>> assert daily_counts.shape == (731, 7)
```
failure_message: daily counts should have 731 rows and 7 columns
```
>>> def stat_summary_check(data):
...     min_val = np.min(data)
...     q25 = np.percentile(data, 25)
...     median = np.median(data)
...     q75 = np.percentile(data, 75)
...     max_val = np.max(data)
...     mean = np.mean(data)
...     std_dev = np.std(data, ddof=1)
...     return tuple((float(round(val, 2)) for val in (min_val, q25, median, q75, max_val, mean, std_dev)))
>>> check_tuple = (2.0, 315.5, 713.0, 1096.0, 3410.0, 848.18, 686.62)
>>> assert np.allclose(stat_summary_check(daily_counts['casual']), check_tuple, 0.001)
```
failure_message: check that you are calculating the sum for the registered column when working with the groupby object
```
>>> def stat_summary_check(data):
...     min_val = np.min(data)
...     q25 = np.percentile(data, 25)
...     median = np.median(data)
...     q75 = np.percentile(data, 75)
...     max_val = np.max(data)
...     mean = np.mean(data)
...     std_dev = np.std(data, ddof=1)
...     return tuple((float(round(val, 2)) for val in (min_val, q25, median, q75, max_val, mean, std_dev)))
>>> check_tuple = (20.0, 2497.0, 3662.0, 4776.5, 6946.0, 3656.17, 1560.26)
>>> assert np.allclose(stat_summary_check(daily_counts['registered']), check_tuple, 0.001)
```
failure_message: check that you are calculating the sum for the casual column when working with the groupby object


Tests for q1c (points: 5)
```
>>> assert ax1c.get_legend()
```
failure_message: Add a legend please
```
>>> assert ax1c.get_title()
```
failure_message: Add a title please
```
>>> assert ax1c.get_xlabel()
```
failure_message: Add an x label please
```
>>> assert ax1c.get_ylabel()
```
failure_message: Add a y label please


Tests for q1e (points: 5)
```
>>> assert len(ax1e.lines) > 0
```
failure_message: add a line of slope 1 intercept 0 to your plot please
```
>>> assert ax1e.get_aspect() == 1
```
failure_message: set your figure aspect ratio to be 1 please
```
>>> assert ax1c.get_title()
```
failure_message: Add a title please
```
>>> assert ax1c.get_xlabel()
```
failure_message: Add an x label please
```
>>> assert ax1c.get_ylabel()
```
failure_message: Add a y label please


Tests for q1f (points: 5)
```
>>> assert len(ax1f.lines) > 0
```
failure_message: Add a line of slope 1 intercept 0 to your plot please
```
>>> assert ax1f.get_legend()
```
failure_message: Add a legend please
```
>>> assert ax1f.get_title()
```
failure_message: Add a title please
```
>>> assert ax1f.get_xlabel()
```
failure_message: Add an x label please
```
>>> assert ax1f.get_ylabel()
```
failure_message: Add a y label please


Tests for q1h (points: 5)
```
>>> assert ax1h.get_legend()
```
failure_message: Add a legend please
```
>>> assert ax1h.get_title()
```
failure_message: Add a title please
```
>>> assert ax1h.get_xlabel()
```
failure_message: Add an x label please
```
>>> assert ax1h.get_ylabel()
```
failure_message: Add a y label please
